## Problem 2 - Temperature classifier (3 points)

In [Problem 2 of Exercise 3 last week](https://github.com/Geo-Python-2022/Exercise-3/blob/master/Exercise-3-problem-2.ipynb) we practiced how to classify temperatures into four different classes: cold, slippery, comfortable, and warm. Let's continue working with the same idea.

Here, you should create a function called `temp_classifier` that accepts a temperature value in Celsius that will be reclassified into integer numbers 0-3 based on following criteria:

| Return value | Classification criteria                                                  |
| :----------: | :----------------------------------------------------------------------- |
| 0            | Temperatures below -2 degrees Celsius                                    |
| 1            | Temperatures equal or warmer than -2, but less than +2 degrees Celsius   |
| 2            | Temperatures equal or warmer than +2, but less than +15 degrees Celsius  |
| 3            | Temperatures equal or warmer than +15 degrees Celsius                    |

### Tips for completing this problem

- Use **exactly** the same variable names as in the instructions because your answers will be automatically graded, and the tests that grade your answers rely on following the same formatting or variable naming as in the instructions.
- **Please do not change the file names**. Do all of your editing in the provided `Exercise-4-problem-2.ipynb` file (this file).
- **Please do not copy/paste code cells in this notebook**. We use an automated grading system that will fail if there are copies of code cells.

### Scores for this problem

**Your score on this problem will be based on following criteria:**

- Creating a new function called `temp_classifier` **(0.5 points)**
- Defining one parameter called `temp_celsius` in your `temp_classifier` function **(0.5 points)**
- Reclassifying input temperatures based on the criteria in the above table and returning the reclassified value as an integer number (i.e., 0, 1, 2, or 3) **(1 point)**
- Adding comments in your code and a docstring that explains how to use your `temp_classifier` function (i.e., you should write the purpose of the function, parameters, and returned values) **(1 point)**
- Uploading your notebook to your GitHub repository for this week's exercise

## Part 1 (2 points)

Create a function in the cell below to classify temperatures as directed.

- Your function should be called `temp_classifier`
- It should have a single parameter called `temp_celsius`
- It should properly classify temperatures for a given temperature input value

<!-- grade_id: problem_2_ag_p1_temp_classifier, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**Use** your newly created function to print the answers to the following questions:

```python
# 1. What is the class value for 16.5 degrees Celsius?

# 2. What is the class value for +2 degrees Celsius?

```

### Problem 2 tests

Run the code cells below to check your answer. Note that you are not able to modify the test cells. Your solution is on the right track if you don't see any errors when running these code cells.

<!-- grade_id: problem_2_ag_p1_temp_classifier_test1, points: 0.5, grade: true, solution: false -->
```python
import inspect

# Check that function exists
assert inspect.isfunction(temp_classifier)
```

<!-- grade_id: problem_2_ag_p1_temp_classifier_test2, points: 0.5, grade: true, solution: false -->
```python
# Check that the function has a single parameter and the pamameter name is correct
params = list(inspect.signature(temp_classifier).parameters.keys())
assert len(params) == 1
assert params[0] == 'temp_celsius'
```

Check that the function produces correct answers for selected values:

<!-- grade_id: problem_2_ag_p1_temp_classifier_test3, points: 0.25, grade: true, solution: false -->
```python
# 1. What is the class value for 16.5 degrees (Celsius)?
assert temp_classifier(16.5) == 3, 'Wrong class'
print("ok :)")
```

<!-- grade_id: problem_2_ag_p1_temp_classifier_test4, points: 0.25, grade: true, solution: false -->
```python
# 2. What is the class value for +2 degrees (Celsius)?
assert temp_classifier(2) == 2, 'Wrong class'
print("ok :)")
```

<!-- grade_id: problem_2_ag_p1_temp_classifier_test5, points: 0.25, grade: true, solution: false -->
```python
# 3. What is the class value for +1 degrees (Celsius)?
assert temp_classifier(1) == 1, 'Wrong class'
print("ok :)")
```

<!-- grade_id: problem_2_ag_p1_temp_classifier_test6, points: 0.25, grade: true, solution: false -->
```python
# 4. What is the class value for -5 degrees (Celsius)?
assert temp_classifier(-5) == 0, 'Wrong class'
print("ok :)")
```

## Part 2 (1 point)

Check your code.

 - Are all the variable names correct?
 - Does your code contain enough comments?
 - Does your function have an informative dosctring?

<!-- grade_id: problem_2_manual_check, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

#### Done!

That's it! Now you are ready to continue with [Problem 3](Exercise-4-problem-3.ipynb).